---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Introduction to Decision Trees in Python

```yaml
type: VideoExercise
key: e5fc3c3bf3
xp: 50
```

`@projector_key`
d9c3eb98666c5131456e5f5e08fa823f
